<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

**Übung 1: Welche Funktion hat der Sekundärspeicher in einem Computer?**

a\) Alle Berechnungen und die Logik des Programms ausführen\
b) Abruf von Webseiten über das Internet\
c) Informationen langfristig zu speichern, auch über einen Stromausfall hinaus\
d) Eingaben des Benutzers entgegennehmen

<!-- grade_id: cell-020b60a969547677, points: 0, grade: false, solution: true -->
```python
antwort = None

# Fügen Sie eine neue Zuweisung ein, in der Sie antwort 
# auf das korrekte Ergebnis a, b, c, oder d setzen

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-edc656af33ac5bfb, points: 1, grade: true, solution: false -->
```python
assert type(antwort) == str, "Aufg. 2: Setzen Sie antwort auf 'a', 'b', 'c', oder 'd'"
```

**Aufgabe 4: Was enthält Maschinencode?**

a\) Der Python-Interpreter\
b) Die Tastatur\
c) Die Python-Quelldatei\
d) Ein Textverarbeitungsdokument

<!-- grade_id: cell-2e7244fa7c982f1a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eedd352e5150629a, points: 1, grade: true, solution: false -->
```python
assert type(antwort) == str, "Aufg. 4: Setzen Sie antwort auf 'a', 'b', 'c', oder 'd'"
```

**Aufgabe 5: Was ist falsch an folgendem Code? Verbessern Sie in in der folgenden Zelle.**

```python
>>> primt 'Hello world!'
File “<stdin>“, line 1
  primt 'Hello world!'
                     ^
SyntaxError: invalid syntax
```

<!-- grade_id: cell-41c9bb1915328d6e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%%capture cap --no-stderr --no-display
c = _i
exec(c)
```

<!-- grade_id: cell-395e4adf2fbca228, points: 1, grade: true, solution: false -->
```python
assert "Hello world!" in cap.stdout, f"Aufg. 5: Die Ausgabe {s} stimmt nicht."
assert "print" in c, "Aufg. 5: Verwenden Sie die print Funktion"
```

**Übung 6: Wo im Computer wird eine Variable wie `x` gespeichert, nachdem die folgende Python-Zeile beendet ist?**

```python
x = 123
```

a\) CPU\
b) Hauptspeicher\
c) Sekundärspeicher\
d) Eingabegeräte\
e) Ausgabegeräte

<!-- grade_id: cell-a5f0e39063fadeb4, points: 0, grade: false, solution: true -->
```python
antwort = None

# Fügen Sie eine neue Zuweisung ein, in der Sie antwort 
# auf das korrekte Ergebnis a, b, c, d, oder e setzen

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-17b0f6d3d903878b, points: 1, grade: true, solution: false -->
```python
assert type(antwort) == str, "Aufg. 6: Setzen Sie antwort auf 'a', 'b', 'c', 'd', oder 'e'"
```

**Übung 7: Was wird das folgende Programm ausgeben:**

```python
x = 43
x = x + 1
print(x)
```

a\) 43\
b) 44\
c) x + 1\
d) Fehler, denn x = x + 1 ist mathematisch nicht möglich

<!-- grade_id: cell-0ae852e7c8b70d30, points: 0, grade: false, solution: true -->
```python
antwort = None

# Fügen Sie eine neue Zuweisung ein, in der Sie antwort 
# auf das korrekte Ergebnis a, b, c, oder d setzen

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9bd2656cdb70fa34, points: 1, grade: true, solution: false -->
```python
assert type(antwort) == str, "Aufg. 7: Setzen Sie antwort auf 'a', 'b', 'c', oder 'd'"
```