<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

**Übung 1: Führen Sie die folgenden Python-Ausdrücke aus, um zu sehen, was sie bewirken:**


```python
5
```

```python
x = 5
```

```python
x + 1
```

**Übung 2: Schreiben Sie ein Programm, das mit `input` einen Benutzer zur Eingabe seines Namens auffordert und ihn dann begrüßt. Der Name soll unter dem Bezeichner `name` gespeichert werden.**


    Gib deinen Namen ein: Heiner
    Hallo Heiner

<!-- grade_id: cell-59651d8333ced2e5, points: 0, grade: false, solution: true -->
```python
name = None
# YOUR CODE HERE
raise NotImplementedError()
```

```python
c = _i
input_old = input

def input(a):
    print(a)
    return "Test"

exec(c)
input = input_old
```

<!-- grade_id: cell-e7ee37613ad722a7, points: 1, grade: true, solution: false -->
```python
assert type(name) == str, "Aufg. 2: Verwenden Sie die Variable name"
```

**Übung 3: Schreiben Sie ein Programm, das den Benutzer nach Arbeitsstunden und Stundensatz fragt, um den Bruttolohn zu berechnen.**

~~~~
Gib die Arbeitsstunden an: 35
Gib den Stundensatz an: 2.75
Zahlung: 96.25
~~~~

<!-- grade_id: cell-5f117b04032cc022, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%%capture cap --no-stderr --no-display
c = _i
input_old = input

def input(a):
    print(a)
    input.counter += 1
    if (input.counter == 1) :
        return "40"
    elif (input.counter == 2) :
        return "12"
    else:
        raise

input.counter = 0

exec(c)
input = input_old
```

<!-- grade_id: cell-b2b3362a4e701fcd, points: 1, grade: true, solution: false -->
```python
assert 'input' in c, "Aufg. 3: Verwenden Sie die input Funktion"
assert '480' in cap.stdout, "Aufg. 3: Die Berechnung stimmt nicht"
```

**Übung 4: Angenommen, wir führen die folgenden Zuweisungen aus:**

```python
breite = 17
hoehe = 12.0
```

Schreiben Sie für jeden der folgenden Ausdrücke den Wert des Ausdrucks und den Datentyp (des Wertes des Ausdrucks).

1.  `breite//2`

2.  `breite/2.0`

3.  `hoehe/3`

4.  `1 + 2 * 5`

```python
# Nutzen Sie diese Zelle zum Testen Ihrer Ergebnisse
```

**Übung 5: Schreiben Sie ein Programm, das den Benutzer zur Eingabe einer Celsius-Temperatur auffordert, die Temperatur in Fahrenheit konvertiert und die konvertierte Temperatur ausgibt.**

**Die Umrechnung von Celsius zu Fahrenheit lautet: °F = °C * 1,8 + 32**

    Gib die Temperatur in Celsius ein: 37
    Fahrenheit in 98.6


<!-- grade_id: cell-31dd2bd8f019cbfb, points: 0, grade: false, solution: true -->
```python
celsius = None
fahrenheit = None
# YOUR CODE HERE
raise NotImplementedError()
print("Fahrenheit:", fahrenheit)
```

```python
%%capture cap --no-stderr --no-display
c = _i
input_old = input

def input(a):
    print(a)
    return "37"

exec(c)
input = input_old
```

<!-- grade_id: cell-e1856df3b60d13a0, points: 1, grade: true, solution: false -->
```python
assert 'input' in c, "Aufg. 5: Verwenden Sie die input Funktion"
assert 'Fahrenheit: 98.60' in cap.stdout, "Aufg. 5: Die Berechnung stimmt nicht"
assert cap.stdout.count('Fahrenheit') == 1, "Aufg. 5: Die Ausgabe stimmt nicht"
```