<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

Bitte führen Sie die folgende Code-Zelle zuerst aus. Dies wird zur Auswertung der Lösungen benötigt.

```python
def myinput(inputs):
    def foo(a):
        print(a)
        input.counter += 1
        if input.counter > len(inputs) :
            raise "Nicht genügend inputs"
        else:
            return inputs[input.counter-1]
    foo.counter = 0
    return foo
```

**Übung 1: Schreiben Sie Ihr Programm zur Lohnberechnung so um, dass ein Mitarbeiter das 1,5-fache des Stundensatzes für Arbeitsstunden über 40 Stunden erhält.**

```python
Gib die Arbeitsstunden an: 45
Gib den Stundensatz an: 10
Zahlung: 475.0
```

<!-- grade_id: cell-97e7f0155847f6ae, points: 0, grade: false, solution: true -->
```python
zahlung = None
# YOUR CODE HERE
raise NotImplementedError()
print("Zahlung: ", zahlung)
```

```python
%%capture cap --no-stderr --no-display
c = _i
input_old = input
input = myinput(['45','10'])
exec(c)
input = input_old
```

<!-- grade_id: cell-7334c106f1cdedc1, points: 1, grade: true, solution: false -->
```python
assert 'input' in c, "Aufg. 1: Verwenden Sie die input Funktion"
assert 'Zahlung:  475.0' in cap.stdout, "Aufg. 1: Die Berechnung stimmt nicht"
assert cap.stdout.count('Zahlung') == 1, "Aufg. 1: Die Ausgabe stimmt nicht"
del cap
```

**Übung 2: Schreiben Sie Ihr Programm zur Lohnberechnung unter Verwendung von `try` und `except` so um, dass Ihr es nichtnumerische Eingaben elegant behandelt, indem es eine Meldung ausgibt und das Programm beendet. Im Folgenden sehen Sie zwei beispielhafte Ausführungen des Programms:**

    Gib die Arbeitsstunden an: 45
    Gib den Stundensatz an: neun
    Fehler, bitte gib eine Zahl ein

    Gib die Arbeitsstunden an: vierzig
    Fehler, bitte gib eine Zahl ein


<!-- grade_id: cell-13440a24a8cadfa3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%%capture cap --no-stderr --no-display
c = _i
input_old = input
input = myinput(['vierzig'])
exec(c)
input = input_old
```

<!-- grade_id: cell-03f5191ce2d949f4, points: 1, grade: true, solution: false -->
```python
assert 'try' in c and 'except' in c, "Aufg. 2: Verwenden Sie try/except"
assert 'Zahlung:' not in cap.stdout, "Aufg. 2: Diese Ausgabe sollte nicht erscheinen"
assert cap.stdout.count('Fehler, bitte gib eine Zahl ein') == 1, "Aufg. 2: Die Ausgabe stimmt nicht"
del cap 
```

**Übung 3: Schreiben Sie ein Programm, das nach einem Wert zwischen 0,0 und 1,0 fragt. Wenn die Punktzahl außerhalb des Bereichs liegt, geben Sie eine Fehlermeldung aus. Wenn die Punktzahl zwischen 0,0 und 1,0 liegt, geben Sie eine Note anhand der folgenden Tabelle aus:**
    
     Punkte   Note 
     >= 0.9     1  
     >= 0.8     2  
     >= 0.7     3  
     >= 0.6     4  
      < 0.6     5  
       
Beispiele:

    Punkte eingeben: 0.95
    1

    Punkte eingeben: perfekt
    Falsche Bitte geben Sie die Punkte von 0.0 bis 1.0 ein

    
    Punkte eingeben: 2
    Falsche Punkanzahl


Führen Sie das Programm wiederholt wie oben gezeigt aus, um die verschiedenen Werte für die Eingabe zu testen.

<!-- grade_id: cell-f26787cc665ecd15, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

```python
%%capture cap --no-stderr --no-display
c = _i
input_old = input
input = myinput(['0.95','perfekt','2'])
exec(c)
exec(c)
exec(c)
input = input_old
```

<!-- grade_id: cell-ae1436d69c7e622c, points: 1, grade: true, solution: false -->
```python
lines = cap.stdout.splitlines()
assert '1' in lines[1], "Aufg. 3: Fehler bei Beispiel 1"
assert 'Bitte geben Sie die Punkte' in lines[3], "Aufg. 3: Fehler bei Beispiel 2"
assert 'Falsche Punktanzahl' in lines[5], "Aufg. 3: Fehler bei Beispiel 3"
del cap
```