<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

**Übung 1 <br>
Es soll ein einfaches Programm geschrieben werden, das die Funktionsweise des Befehls `grep` unter Unix simuliert. Der Benutzer soll aufgefordert werden, einen regulären Ausdruck einzugeben, mit dem dann die Anzahl der Zeilen in einer Datei (Parameter `filename`) gezählt wird, die mit dem regulären Ausdruck übereinstimmen:** 

*$ python grep.py<br>
Gib eine RegEx an: ^Author <br>
mbox.txt hat 1798 Zeilen die auf den Ausdruck "^Author" passen*

**Geben Sie die Anzahl der gezählten Zeilen zurück.**

<!-- grade_id: cell-5589f54f815efd9c, points: 0, grade: false, solution: true -->
```python
import re
def grep_simulation(filename):
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d51d5d317e82c222, points: 1, grade: true, solution: false -->
```python
"""grep_simulation Tests"""
from nose.tools import assert_equal

def myinput(inputs):
    def foo(a):
        print(a)
        input.counter += 1
        if input.counter > len(inputs) :
            raise "Nicht genügend inputs"
        else:
            return inputs[input.counter-1]
    foo.counter = 0
    return foo

from nose.tools import assert_equal

fname = "files/mbox-short.mbox"

input_old = input
input = myinput(["^Author"])
r1 = grep_simulation(fname)
input = myinput(["^X-"])
r2 = grep_simulation(fname)
input = input_old

assert 27 == r1, "Der Regex wurde nicht korrekt ausgewertet"
assert 216 == r2, "Der Regex wurde nicht korrekt ausgewertet"


```

**Übung 2 <br>
Es soll ein Programm geschrieben werden, das nach Zeilen der folgenden Form sucht:**

`New Revision: 39772`

**Die Zahl soll aus jeder der Zeilen mit einem regulären Ausdruck und der Methode `findall()` extrahiert werden. Es soll dann der Durchschnitt der Zahlen berechnet und als Ganzzahl zurückgegeben werden.**


<!-- grade_id: cell-39c2a2a83e599bad, points: 0, grade: false, solution: true -->
```python
import re
def find_by_regex(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1fc4221d62a6df92, points: 1, grade: true, solution: false -->
```python
"""find_by_regex tests"""
from nose.tools import assert_equal

fname = "files/mbox-short.mbox"
res = find_by_regex(fname)
assert type(res) == int, "Die Funktion soll ein int-Objekt zurückliefern"
assert_equal(39756, res)
```