<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

In diesem Aufgabenblatt soll es um das Thema **NumPy** gehen.
Wir benötigen für die Aufgaben hauptsächlich zwei Bibliotheken.
Einerseits `numpy`, das wir wie üblich mit `import numpy as np` importieren.
Wir können nun auf alles, was "in NumPy steckt" über den *Namensraum* `np` zugreifen.
Die zweite Bibliothek, die wir verwenden, ist die *Matplotlib*, aus der wir das Modul `pyplot` unter dem Namensraum `plt` importieren.
Wie Matplotlib funktioniert, werden Sie zum kommenden Termin anschauen.
Wir verwenden hier aber bereits einige Funktionen, um unsere Daten zu visualisieren.

```python
import matplotlib.pyplot as plt
import numpy as np
```

### Matrizenmultiplikation
Zum Einstieg in das Thema **NumPy** wollen wir einige `ndarrays` anlegen und Berechnungen durchführen.

**Aufgabe:** Erstellen Sie die folgenden Matrizen und Vektoren:

$
A=\begin{bmatrix} 1 & 2 & -3\\ 3 & 4 & -1 \end{bmatrix},\ \ 
B=\begin{bmatrix} 1 & -5 & 1\\ 1 & 4 & 5\\ 2 & -1 & 6 \end{bmatrix},\ \ 
y=\begin{bmatrix} 2\\ -4\\ 1\end{bmatrix},\ \ 
z=\begin{bmatrix} -15\\ -8\\ -22\end{bmatrix}
$

und führen Sie folgende Berechnungen aus:

- $a = AB$ (Matrizenmultiplikation)<br>
- $b = BA^T$<br>
- $c = By$<br>
- $d = y^Tz$ (Skalarprodukt, engl. *inner product* oder*dot-product*)<br>
- $e = yz^T$ (Äußeres Produkt, engl. *outer product*)<br>

<!-- grade_id: cell-458a26a7ae453abd, points: 0, grade: false, solution: true -->
```python
A = B = y = z = None
a = b = c = d = e = None

# YOUR CODE HERE
raise NotImplementedError()

print("\nA = \n", A)
print("\nB = \n", B)
print("\ny = \n", y)
print("\nz = \n", z)

print("\nAB = \n", a)
print("\nBA' = \n", b)
print("\nBy = \n", c)
print("\ny'z = \n", d)
print("\nyz' = \n", e)
```

<!-- grade_id: cell-6bc2c65eeee47e7e, points: 0, grade: true, solution: false -->
```python
# Testzelle
assert a.sum() == 20
assert b.sum() == -44
assert c.sum() == 28
assert d.sum() == -20
assert e.sum() == 45
```

### Zufallszahlen

Wir wollen nun ein 3-dimensionale *ndarray* anlegen und dieses mit Zufallszahlen füllen.

Um zufällige Werte zu generieren, können wir die Funktion `random` aus dem Modul `np.random` verwenden (ja, die Funktion und das Modul heißen beide *random*). Diese Funktion erzeugt ein mit Zufallswerten aus dem Bereich $[0,1)$ gefülltes *ndarray* in der Größe, die Sie als ersten Parameter angeben. Dieser Parameter kann ein Tupel sein, z.B. steht `(3,3,3)` für eine Größe von $3\times{}3\times{}3$, also einer 3-dimensionalen Matrix mit 27 Werten.

Es gibt Zufallszahlengeneratoren für verschiedene Zufallsverteilungen. Nach welcher Verteilung sind die Zahlen von `random` verteilt? Sie können zur Beantwortung der Frage in die Dokumentation schauen. Einen guten Eindruck über die Verteilung erhält man aber auch, indem man sich den Mittelwert der Daten ansieht und berechnet, wie viele Zahlen in bestimmte Werte-Intervalle fallen.

**Aufgabe:** Legen Sie mit `np.random.random` ein 3D *ndarray* der Größe $100\times{}100\times{}10$ an. Beantworten Sie folgende Fragen:

 a) Wie viele der Werte sind kleiner als $0.1$?<br>
 b) Wie viele der Werte  sind größer als $0.1$ und kleiner als $0.2$?<br>
 c) Was ist das arithmetische Mittel der Werte?<br>

<!-- grade_id: cell-c1200e9fa024c2d8, points: 0, grade: true, solution: true -->
```python
a = None # Werte kleiner 0.1
b = None # Werte im Intervall (0.1,0.2)
c = None # Mittelwert

# YOUR CODE HERE
raise NotImplementedError()
a, b, c
```

Mit der Matplotlib Funktion `pyplot.imshow()` können NumPy Arrays als Bilder angezeigt werden.
Dabei sind die ersten beiden Dimensionen die Zeilen bzw. die Spalten des Bildes.
Besitzt das Array nur 2 Dimensionen, so werden die Werte als *Helligkeitswerte* auf einer Farbskala interpretiert. 
Standardmäßig ist die Blau-Grün-Gelbe Farbskala `viridis` voreingestellt.
Falls Sie eine andere Skala nutzen wollen, können Sie den Parameter `cmap` verwenden, z.B. mit `cmap="gray"` für ein Graustufenbild.

Besitzt das NumPy Array 3 Dimensionen, können auch Farbbilder angezeigt werden.
Bei einem 3-Tupel in der 3. Dimension wird dieses Tupel als *RGB-Pixel* interpretiert, bei einem 4-Tupel steht der letzte Wert für die Transparenz des Pixels.


```python
# Ein 3D Array der Größe 200x200x3 mit 120.000
# Zufallswerten im Bereich [0,255) anlegen
rand_array = 255 * np.random.random((200,200,3))

# Standardmäßig verwendet NumPy float-Werte unser Pixel-Werte sind
# 8-bit vorzeichenlose Ganzzahlen (uint8)
rand_array = rand_array.astype("uint8")

# Die 1. und 2. Dimension sind die Zeilen und
# Spalten des Bildes. Die drei werte der 3. Dimension
# sind die RGB-Farbanteile des entsprechenden Pixels
plt.imshow(rand_array)
```

**Aufgabe:** Erzeugen Sie ein Bild mit der Schweizer Flagge. Legen Sie dazu ein Bild mit 300 mal 300 Pixeln an. Der Hintergrund soll Rot sein, in der Mitte des Bildes soll ein weißes Kreuz erscheinen.

<!-- grade_id: cell-a0940d8baeb63c65, points: 0, grade: true, solution: true -->
```python
imgsize = (300,300,3)
swiss = np.zeros(imgsize, dtype="uint8")

# Ein weißes Pixel:
weiss = (255,255,255)

# YOUR CODE HERE
raise NotImplementedError()

plt.imshow(swiss)
```

**Aufgabe (optional):** Erzeugen Sie ein $8\times{}8$ Graustufen-Bild mit einem Schachbrett-Muster.

<!-- grade_id: cell-bd085c84b461b1a0, points: 0, grade: true, solution: true -->
```python
checkboard = np.zeros((8,8), dtype="uint8")

# YOUR CODE HERE
raise NotImplementedError()

plt.imshow(checkboard, cmap="gray")
```

### Gesichter erkennen

Der *Labeled Faces in the Wild* Datensatz (LFW) ist eine Sammlung von Portrait-Fotos, die sich als Beispiel-Daten für Problemstellungen wie Gesichtserkennung einsetzen lassen.
Die Datenbank wurde an der University of Massachusetts in Amherst erstellt und wird dort auch weiterhin gepflegt.
Sie enthält über 13.000 Bilder von mehr als 5.000 Personen.

Über die *Scikit-Learn* Bibliothek können wir den Datensatz recht einfach laden.
In der Code-Zelle unten, laden wir die Bilder von denjenigen Personen herunter, die mit 100 oder mehr bildern im Datensatz vertreten sind.

Das Ergebnis ist eine Art Dictionary mit folgenden Schlüsseln
- `data` Die Graustufen-Pixel aller geladenen Bilder als eindimensionales NumPy-Array. Jedes Bild besitz 62x47 Pixel und damit 2914 Graustufenwerte
- `images` Die gleichen Werte wie in `data` nur als zweidimensionale NumPy-Arrays der Größe 62x47
- `target` Die ID der Person auf dem Bild. Das Attribut `target_names` liefert ein Array mit den Namen zu den entsprechenden IDs zurück
- `DESCR` Eine Beschreibung zum LFW Datensatz


```python
from sklearn.datasets import fetch_lfw_people

# Lade alle Personen mit mehr als 100 Bildern in der DB herunter
people = fetch_lfw_people(min_faces_per_person=100)
```

Wir sehen, dass es nur wenige Personen mit 100 oder mehr Bildern im Datensatz gibt. Wir können nun nachschauen, um welche Personen es sich handelt:

```python
ids = set(people['target'])
for i in ids:
    print(f"ID{i} -> {people.target_names[i]}")
```

Nun wollen wir ein zufälliges Bild ausgeben.
Dazu generieren wir uns eine Zufallszahl im Bereich von 0 bis zur Anzahl der Bilder im Datensatz.
Wir laden die Bilddaten aus dem Datensatz und geben Sie mit der Funktion `imshow` aus.

**Aufgabe:** Überlegen Sie sich, warum wir hier die Methode `squeeze()` einsetzen. Was tut die Funktion?

```python
num_images, _ = people.data.shape
rand_num = np.random.randint(num_images, size=1)
rand_img = people.images[rand_num].squeeze()
plt.imshow(rand_img, cmap='gray')

```

**Aufgabe:** Schreiben Sie nun eine Funktion `get_images_of(person_id)` die als Ergebnis alle Bilder (`images`) einer bestimmten Person zurückliefert. Die Person soll über die ID, und nicht Ihren Namen ausgewählt werden. Gerhard Schröder z.B. hat die ID 3 im Datensatz. 

<!-- grade_id: cell-0000b26c25e08467, points: 0, grade: true, solution: true -->
```python
def get_images_of(person_id):
    # YOUR CODE HERE
    raise NotImplementedError()

schroeder = get_images_of(3)
plt.imshow(schroeder[0], cmap='gray')
```

**Aufgabe:** Schreiben Sie eine Funktion `avg_image(person_id)`, die das "Durchschnittsbild" einer Person berechnet.
Dieses Durchschnittsbild erhalten Sie, indem Sie die Pixel aller Bilder (an jeweils einer Stelle) aufsummieren und davon den Mittelwert wilden.

Beachten Sie, dass die einzelnen Bilder 2-dimensionale NumPy Arrays sind, also Matrizen.
Sie können die Werte zweier Matrizen punktweise addieren, indem Sie einfach den `+`-Operator verwenden.
Den Mittelwert erhalten Sie, indem Sie einfach durch die Anzahl der aufaddierten Bilder teilen.

<!-- grade_id: cell-c6a036b19e61a380, points: 0, grade: true, solution: true -->
```python
def avg_image(person_id):
    # YOUR CODE HERE
    raise NotImplementedError()

avg_schroeder = avg_image(3)

plt.imshow(avg_schroeder, cmap="gray")
```

**Aufgabe:** Schreiben Sie eine Funktion `img_diff(i0, i1)`, die die mittlere Abweichung der Pixel-Werte zweier Graustufenbilder `i0` und `i1` zurückliefert.

Sie können die Aufgabe lösen, indem Sie die Werte der beiden Bilder voneinander abziehen und dann mit `np.abs()` die absoluten Werte bestimmen. `np.sum()` gibt Ihnen die Summer aller Werte in einem `ndarray` zurück, das Attribut `size` liefert Ihnen die Anzahl der Werte im Array.

<!-- grade_id: cell-c58568c51a9975c9, points: 0, grade: true, solution: true -->
```python
def img_diff(i0, i1):
    # YOUR CODE HERE
    raise NotImplementedError()


bush = get_images_of(2)
avg_bush = avg_image(2)

img_diff(avg_bush, bush[0])

```

Wir können nun versuchen, mit unseren Durchschnittsbildern eine Personenerkennung zu Programmieren.
Die Idee ist folgende:
- Wir kennen die Durchschnittsbildern für verschiedene Personen
- Wir bekommen ein neues Bild und sollen schätzen, welche Person darauf abgebildet ist
- Wir nehmen an, dass das Durchschnittsbild für die gesuchte Person den geringsten Fehler (geringsten Unterschied der Pixel Werte) zum vorliegenden Bild besitzt


**Aufgabe:** Berechnen Sie die Durchschnittsbilder von Gerhard Schröder und George W. Bush. Durchlaufen Sie dann alle Bilder von George W. Bush und bestimmen Sie jeweils, ob die Abweichung des Bildes zum Durchschnittsbild von Bush oder zum Durchschnittsbild von Schröder geringer ist. In wie vielen Fällen ist die Abweichung zum Durchschnittsbilde von George W. Bush geringer (*richtige* Vorhersage), in wie vielen Fällen zu dem von Gerhard Schröder (*falsche* Vorhersage)?

<!-- grade_id: cell-e28d73704047c5c8, points: 0, grade: true, solution: true -->
```python
bush = get_images_of(2)
avg_bush = avg_image(2)
avg_schroeder = avg_image(3)

richtig = 0
falsch = 0
# YOUR CODE HERE
raise NotImplementedError()
print("richtig: ", richtig)
print("falsch: ", falsch)

```