<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

## Matrix-Klasse

In diesem Arbeitsblatt soll eine Klasse für mathematische Matrizen erstellt werden.
Die Klasse `Matrix` besitzt ein Klassen-Attribut `klassenname` und soll folgende Instanz-Attribute erhalten:
- `zeilen`: Die Anzahl der Zeilen in der Matrix
- `spalten`: Die Anzahl der Spalten in der Matrix
- `data`: Werte der Matrix gespeichert als zweidimensionale Liste

In der folgenden Code-Zelle sind ebenfalls die Methoden `__str__` und `__repr__` implementiert.
Dabei handelt es um sogenannte *Magic Methods*, das sind Methoden mit einem festem Namen und fester Bedeutung.
Damit sie sich direkt von anderen Methoden abheben, beginnen und enden die Bezeichner von *Magic Methods* immer mit zwei Unterstrichen.
*Magic Methods* werden zu bestimmten Ereignissen aufgerufen.
Wir können diese *Magic Methods* (und *Magic Attributes*) überschreiben, um die vordefinierte Funktionalität zu ersetzen oder zu erweitern.

Die `__init__(self,...)` Funktion ist ein Beispiel für eine *Magic Method*. Wenn Sie eine neue Instanz einer Klasse erzeugen, sucht der Interpreter nach einer Methode mit genau diesem Namen und führt sie aus.
Die *Magic Methods* `__str__` und `__repr__` werden aufgerufen, wenn ein Objekt zu einem String umgewandelt wird, bzw. wenn die Darstellung des Objekts auf der Kommandozeile ausgegeben wird.

#### Aufgabe: Konstruktor hinzufügen
Programmieren Sie den Konstruktor der Klasse `Matrix`.
Der Konstruktor soll mit mindestens zwei Argumenten aufgerufen werden, mit denen die Anzah der Zeilen und die Anzahl der Spalten gesetzt werden können. Das Matrix-Objekt soll die Anzahl der Zeilen und Spalten in den Attributen `zeilen` bzw. `spalten` abspeichern.
Die Werte der Matrix sollen als **eindimensionale Liste** übergeben werden.
Intern sollen die Daten der Matrix als **zweidimensionale Liste** (*Liste von Listen*) unter dem Attributnamen `data` abgespeichert werden.
Data enthält also für jede Zeile der Matrix eine Liste, die die Werte der entsprechenden Zeile enthält.

```python
werte = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11]
m = Matrix(3, 4, werte)
print(m)
```
*Erwartete Ausgabe:*
```
|  0.000  1.000  2.000  3.000 |
|  4.000  5.000  6.000  7.000 |
|  8.000  9.000 10.000 11.000 |
```

##### Erweiterung:

Der Konstruktor soll auch funktionieren, wenn keine Liste von Werten angegeben wird. In diesem Fall soll der Konstruktor ein Matrix mit Nullen erzeugen;

```python
m = Matrix(2, 2)
print(m)
```
*Erwartete Ausgabe:*
```
|  0.000  0.000 |
|  0.000  0.000 |
```


<!-- grade_id: cell-9f6032b8d18c809e, points: 0, grade: false, solution: true -->
```python
from random import random
from random import seed

seed(42)


class Matrix:
    
    klassenname = "Matrix"
    
    # Hier soll der Konstruktor der Klasse angegeben werden, also
    # def __init__(...)
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # __str__() wird aufgerufen, wenn ein Objekt in eine Zeichenkette
    # umgewandelt werden muss. Dies ist z.B. bei 'print' der Fall
    def __str__(self):
        res = ""
        for i in self.data:
            res += '| '
            for j in i:
                res += f"{j:6.3f} "
            res += '|\n'
        return res
    
    # __repr__() soll eine möglichst eindeutige Repräsentation des Objekts
    # als Zeichenkette liefer. Diese Ausgabe erscheint auf der Kommandozeile
    # als Ausgabe für das Objekt
    def __repr__(self):
        return f"\n{self.klassenname}-Objekt ({id(self)}):\n{self.__str__()}"
```

<!-- grade_id: cell-bc3ee71e85698b81, points: 0, grade: true, solution: false -->
```python
# Test Kontruktor mit Werten
row = 3
col = 4
werte = [i for i in range(row*col)]
print("Werte:", werte, '\n')
m0 = Matrix(row, col, werte)
assert m0.data[1][1]==5
print(m0)
```

<!-- grade_id: cell-fa84302a35c4c775, points: 0, grade: true, solution: false -->
```python
# Test Kontruktor ohne Werte
m1 = Matrix(2, 2)
m1
```

#### Aufgabe: Eigene Methode definieren

Die Klasse `Matrix` speichert die Werte der Matrix als zweidimensionale Liste.
Für spätere Aufgaben ist es hilfreich, wenn wir die Werte als eindimensionale Liste auslesen können.
Fügen Sie der Klasse `Matrix` eine neue Methode `get_values` hinzu, die alle Werte der Matrix **zeilenweise** in eine Liste überträgt und diese Liste als Ergebnis zurückliefert.

**Hinweis:** Damit wir die Klasse `Matrix` in mehreren Code-Zellen entwicklen können, müssen wir die ursprüngliche Klasse `Matrix` erweitern.
Dazu erzeugen wir jeweils eine neue Klasse Matrix und *erben* alles aus der zuvor definierten Klasse `Matrix`. Das geht, indem wir bei der Klassen-Definition die Vater-Klasse `Matrix` angeben:
```python
class Matrix(Matrix):
    pass
```

<!-- grade_id: cell-0ee3604236607f9c, points: 0, grade: false, solution: true -->
```python
class Matrix(Matrix):
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-869e0f09a57a92a4, points: 0, grade: true, solution: false -->
```python
# Test get_values
werte = [i for i in range(row*col)]
m0 = Matrix(row, col, werte)
res = m0.get_values()
assert m0.get_values()==werte
print(res)
```

#### Aufgabe: Matrix transponieren

Eine der wichtigsten Operationen in der linearen Algebra ist das Transponieren einer Matrix.
Die transponierte Matrix erhält man, indem man die Zeilen und Spalten der ursprünglichen Matrix vertauscht.
Aus einer $3\times{}5$ Matrix, beispielsweise, wird so eine $5\times{}3$ Matrix.

Ergänzen Sie in der Klasse `Matrix` eine Methode `transpose` die für ein Matrix-Objekt die transponierte Matrix als **neue** Matrix erzeugt und zurückliefert. Die Werte der ursprünglichen Matrix sollen also nicht verändert werden.

**Beispiel**<br>
*Ursprüngliche Matrix:*
```
|  0.000  1.000  2.000  3.000 |
|  4.000  5.000  6.000  7.000 |
|  8.000  9.000 10.000 11.000 |
```
*Transponierte Matrix:*
```
|  0.000  4.000  8.000 |
|  1.000  5.000  9.000 |
|  2.000  6.000 10.000 |
|  3.000  7.000 11.000 |
```

**Hinweis:** In der Code Zelle wird ein sogenannter *Property Decorator* verwendet, um der Methode `T` eine spezielle Bedeutung zu geben. Wenn wir der Definition von `T` ein `@property` vorschalten, kann man `T` wie ein Attribut der Klasse Matrix verwenden. Diese Kurzschreibweise ist hilfreich für oft verwendete Funktionen.

<!-- grade_id: cell-e882b14706c2c703, points: 0, grade: false, solution: true -->
```python
class Matrix(Matrix):
    # Hier soll die Methode 'transpose' implementiert werden
    # YOUR CODE HERE
    raise NotImplementedError()
    
    @property
    def T(self):
        return self.transpose()
```

<!-- grade_id: cell-c59b60929b54ad65, points: 0, grade: true, solution: false -->
```python
# Test Transpose
werte = [i for i in range(row*col)]
m0 = Matrix(row, col, werte)
m1 = m0.T
for i in range(m0.zeilen):
    for j in range(m0.spalten):
        assert m0.data[i][j] == m1.data[j][i]
```

#### Aufgabe: Matrizen addieren

Erweitern Sie die Klasse `Matrix` um eine Methode `add` mit der zwei Matrizen gleichen Typs (also mit jeweils gleicher Anzahl von Zeilen und Spalten) addiert und das Ergebnis als neue Matrix zurückliefert. Die Methode soll prüfen, ob der zweite Operand (das $B$ in der Abbildung) ein `Matrix`-Objekt ist. Dafür können Sie z.B. die die Funktion `isinstance` verwenden:

```python
>>> isinstance(42, int)
True
```

Wenn der zweiter Operand keinen gültigen Typ besitzt (also keine `Matrix` ist, oder nach der Erweiterung unten auch keine `int`oder `float` ist), soll die Methode einen `TypeError` erzeugen

<table><tr><td>
<img src="https://upload.wikimedia.org/wikipedia/commons/thumb/5/59/Matrix_addition_qtl1.svg/320px-Matrix_addition_qtl1.svg.png" alt>
</td><td>
<img src="https://upload.wikimedia.org/wikipedia/commons/thumb/4/47/Matrix_addition_qtl2.svg/320px-Matrix_addition_qtl2.svg.png" alt>
</td></tr><tr><td colspan=2, style='text-align:center'>
<b>Matrizenaddition</b> <i>Quelle: Quartl, wikimedia.org, CC BY-SA 3.0</i>
</td></tr></table>



**Erweiterung:**
Erweitern Sie die Methode `add` so, dass man eine Matrix auch mit einer Konstanten addieren kann. Diese Rechenoperation ist in der linearen Algebra zwar nicht definiert, in Anwendungen aber recht praktisch. Wenn also der zweite Operand kein `Matrix`-Objekt ist, sondern ein `int`- oder `float`-Typ, dann soll zu jedem Element in der Matrix dieser Wert addiert werden.<br>
(Wir werden dieses Konzept später als *Broadcasting* kennen lernen)


<!-- grade_id: cell-ff479d98337ffc98, points: 0, grade: false, solution: true -->
```python
class Matrix(Matrix):
    def add(self, other):
        # YOUR CODE HERE
        raise NotImplementedError()
```

<!-- grade_id: cell-8f288cccae185914, points: 0, grade: true, solution: false -->
```python
# Test Add
werte = [i for i in range(row*col)]
m0 = Matrix(row, col, werte)
m1 = Matrix(row, col, werte)
m2 = m0.add(m1)
for i in range(m2.zeilen):
    for j in range(m2.spalten):
        assert m2.data[i][j] == 2*m0.data[i][j]
try:
    m4 = m0.add('Hello')
    assert False, "Solch eine Addition funktioniert nicht"
except TypeError:
    pass
except:
    assert False, "Falscher Fehler-Typ"
    
```

Wenn Sie die Erweiterung implementiert haben, sollte auch folgendes funktionieren:

```python
m0 = Matrix(row, col, [i for i in range(row*col)])
m0, m0.add(1)
```

#### Operatoren überladen

Die *Magic Methods* in Python haben noch eine weitere Funktion, mit ihnen können die bekannten Operatoren **überladen** werden.
Das bedeutet, wir können z.B. selbst definieren, was der *Plus-Operator* `+` für unsere Klasse für eine Bedeutung hat.
Wir müssen dafür einfach die Methode `__add__` definieren.
`__add__` muss neben dem `Matrix`-Objekt `self` noch ein weiteres Objekt erhalten und unbedingt ein neues Objekt zurückliefern.
Das passt genau zu unserer `add`-Methode.
Wir können also einfach `add` in `__add__`aufrufen:

```python
class Matrix(Matrix):   
    
    def __add__(self, other):
        return self.add(other)

```

Nun überprüfen wir, ob das Ergebnis passt:

```python
m0 = Matrix(row, col, [random() for i in range(row*col)])
m1 = Matrix(row, col, [random() for i in range(row*col)])

m0, m1, m0+m1
```

#### Aufgabe: Skalarmultiplikation

Matrizen können mit Skalaren multipliziert werden, die Oeration nennt man entsprechend *Skalarmultiplikation* (übrigens nicht zu verwechseln mit dem *Skalarprodukt*).

${\begin{pmatrix}1&2\\3&4\end{pmatrix}}\cdot 3={\begin{pmatrix}1\cdot 3&2\cdot 3\\3\cdot 3&4\cdot 3\end{pmatrix}}={\begin{pmatrix}3&6\\9&12\end{pmatrix}}$

Eine punktweise Multiplikation zweier Matrizen (also wie bei der Matrizenaddition, nur dass die einzelnen Elmente nicht addiert sondern multipliziert werden) gibt es übrigens auch. Dies ist aber <u>nicht die Matrizenmultiplikation</u>! Man nennt dieses spezielle Produkt auch **Hadamard Produkt**.

<table><tr><td>
<img src="https://upload.wikimedia.org/wikipedia/commons/thumb/0/00/Hadamard_product_qtl1.svg/320px-Hadamard_product_qtl1.svg.png" alt>
</td><td>
<img src="https://upload.wikimedia.org/wikipedia/commons/thumb/b/b9/Hadamard_product_qtl2.svg/320px-Hadamard_product_qtl2.svg.png" alt>
</td></tr><tr><td colspan=2, style='text-align:center'>
<b>Hadamard Produkt</b> <i>Quelle: Quartl, wikimedia.org, CC BY-SA 3.0</i>
</td></tr></table>

**Aufgabe:** Schreiben Sie eine Methode `mul` mit der die Skalarmultiplikation einer Matrix mit einem Skalar sowie das Hadamard-Produkt zweier typgleicher Matrizen berechnet werden kann. Wie bei der Methode `add` soll eine neue Matrix erzeugt werden und im Fall eines flalschen Operanden ein `TypeError` erzeugt werden.

Definieren Sie analog zur Methode `__add__` die Methode `__mul__`, sodass `Matrix`-Objekte nun auch mit dem `*`-Operator multipliziert werden können.


<!-- grade_id: cell-e510fa5014f0fcdb, points: 0, grade: false, solution: true -->
```python
class Matrix(Matrix):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-ee6cf8e9ce35c10e, points: 0, grade: true, solution: false -->
```python
# Test mul
werte = [i for i in range(row*col)]
m0 = Matrix(row, col, werte)
m1 = Matrix(row, col, werte)
m2 = m0.mul(m1)
for i in range(m2.zeilen):
    for j in range(m2.spalten):
        assert m2.data[i][j] == m0.data[i][j]**2
```

Für unsere `Matrix`-Klasse können wir noch viele weitere *Magic Methods* ausimplementieren.
Z.B. die Methoden `__eq__`, um Matrizen auf Gleichheit zu überprüfen oder `__len__` um die Anzahl der Elemente in der Matrix zu bestimmen.
Ebenfalls hilfreich sind die Methoden `__getitem__` und `__setitem__` mit denen man auf einzelne Elemente der Matrix mithilfe der der Index-Notation in `[]` lesen bzw. schreiben kann.

**Aufgabe:** Definieren Sie in der folgenden Code Zelle ebenfalls die Methode `__eq__`, die aufgerufen wird, wenn zwei Matrizen mit `==` auf Gleichheit überprüft werden.

<!-- grade_id: cell-b266e911108dd9e1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
m0 = Matrix(row, col, [i for i in range(row*col)])
assert len(m0)==row*col, "Die __len__-Methode funktioniert nicht"
assert m0[2][2]==10
m0[2][2]=123
assert m0[2][2]==123

m0 = Matrix(row, col, [i for i in range(row*col)])
m1 = Matrix(row, col, [i for i in range(row*col)])
assert m0==m1
m0[1][1]=42
assert m0!=m1
```

#### Aufgabe: Matrizenmultiplikation

Eine der häufigsten Aufgaben im Umgang mit Matrizen ist die Matrizenmultiplikation.
Sie werden in kommenden Modulen noch sehen, das etwa *Neuronale Netze* einen Großteil (>90%) ihrer Berechnungszeit mit Matrizenmultiplikationen berechnen.

Im Unterschied zum elementweisen (Hadamard) Produkt berechnet die Matrizenmultiplikation für jeden Eintrag der Ergebnis-Matrix das Skalarprodukt der entsprechenden Zeile der linken Operanden mit der entsprechenden Spalte des rechten Operanden.

<table><tr><td>
<img src="https://upload.wikimedia.org/wikipedia/commons/thumb/1/18/Matrix_multiplication_qtl1.svg/320px-Matrix_multiplication_qtl1.svg.png" alt>
</td><td>
<img src="https://upload.wikimedia.org/wikipedia/commons/thumb/8/87/Matrix_multiplication_qtl2.svg/320px-Matrix_multiplication_qtl2.svg.png" alt>
</td></tr><tr><td colspan=2, style='text-align:center'>
<b>Matrizenmultiplikation</b> <i>Quelle: Quartl, wikimedia.org, CC BY-SA 3.0</i>
</td></tr></table>

**Aufgabe:** Schreiben Sie eine Methode `matrix_multiply`, die die Matrizenmultiplikation zweier `Matrix`-Objekte berechnet. Die Methode soll eine neue Matrix zurückliefern.

Wir verwenden diese Methode in der *Magic Method* `__matmul__` die immer dann aufgerufen wird, wenn der Matrizenmultiplikatios-Operator `@` verwendet wird.

<!--![](https://github.com/fh-swf-hgi/intro-ml/raw/main/Praktika/A03/matmult.png)-->

<!-- grade_id: cell-5177dd3f3abcfa27, points: 0, grade: false, solution: true -->
```python
class Matrix(Matrix):
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    def __matmul__(self, other):
        return self.matrix_multiply(other)
```

```python
m0 = Matrix(2, 3, [i for i in range(6)])
m1 = Matrix(3, 2, [i for i in range(6)])
m2 = Matrix(2, 2, [10,13,28,40])
assert m0@m1==m2, "Die Matrizenmultiplikation stimmt nicht"

print("m0:\n", m0, sep='')
print("m1:\n", m1, sep='')
print("m0@m1:\n", m0@m1, sep='')

```