<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

**Aufgabe 1:** Shakespeare hat über 20,000 verschiedene Wörter in seinen
Werken verwendet. Aber wie kann man das feststellen? Wie würden Sie die Liste aller Wörter erstellen, die Shakespeare verwendet hat? Würden Sie sein gesamtes Werk herunterladen, es lesen und alle einzigartigen Wörter von Hand nachverfolgen?

Lassen Sie uns stattdessen Python verwenden, um das zu erreichen. Listen Sie alle eindeutigen Wörter, sortiert in alphabetischer Reihenfolge, auf, die in einer Datei romeo.txt gespeichert sind, die eine Teilmenge von Shakespeares Werk enthält.

Erstellen Sie eine Liste mit eindeutigen Wörtern. Schreiben Sie dazu eine Funktion `sort_words`, um die Datei `romeo.txt` zu öffnen und Zeile für Zeile zu lesen. Teilen Sie jede Zeile der Datei mittels der Funktion `split` in eine Liste von Wörtern auf. Prüfen Sie für jedes Wort, ob das Wort bereits in der Liste der eindeutigen Wörter enthalten ist. Wenn das Wort nicht in der Liste der eindeutigen Wörter enthalten ist, fügen Sie es der Liste hinzu. Wenn das Programm mit dem Lesen der Datei fertig, sortieren Sie die Liste der eindeutigen Wörter in alphabetischer Reihenfolge und geben Sie diese als Ergebnis der Funktion `sort_words` (also mit `return`) zurück.

*Tipp: \n gilt als ein Char.*

<!-- grade_id: cell-d343a0656d5e1a2d, points: 0, grade: false, solution: true -->
```python
def sort_words(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-14821efd6cc5f2ba, points: 1, grade: true, solution: false -->
```python
# Test für die Funktion sort_words
from nose.tools import assert_equal, assert_not_equal

fname = "files/romeo.txt"
assert_not_equal(['Arise', 'But', 'It', 'Juliet', 'Who', 'already', 'and', 'breaks\n', 'east', 'envious', 'fair', 'grief', 'is', 'kill', 'light', 'moon\n', 'pale', 'sick', 'soft', 'sun', 'sun\n', 'the', 'through', 'what', 'window', 'with', 'yonder'], sort_words(fname), "Prüfe ob du das Umbruchzeichen entfernst")
assert_equal(['Arise', 'But', 'It', 'Juliet', 'Who', 'already', 'and', 'breaks', 'east', 'envious', 'fair', 'grief', 'is', 'kill', 'light', 'moon', 'pale', 'sick', 'soft', 'sun', 'the', 'through', 'what', 'window', 'with', 'yonder'], sort_words(fname), "Da stimmt etwas noch nicht")
```

**Aufgabe 2:** MBOX (Mailbox) ist ein beliebtes Dateiformat zum Speichern und Freigeben einer Sammlung von E-Mails. Dies wurde von frühen E-MailServern und Desktop-Apps verwendet.

Ohne zu sehr ins Detail zu gehen, ist MBOX eine Textdatei, die E-Mails fortlaufend speichert.
E-Mails werden durch eine spezielle Zeile getrennt, die mit <code>From&nbsp;</code> beginnt (beachten Sie das Leerzeichen).

Wichtig ist, dass Zeilen, die mit `From:` beginnen (beachten Sie den Doppelpunkt), die E-Mail selbst beschreiben und nicht als Trennzeichen fungieren.
Stellen Sie sich vor, Sie schreiben eine minimalistische E-Mail-App, die die E-Mails der Absender im Posteingang des Benutzers auflistet und die Anzahl der E-Mails zählt.

Schreiben Sie ein Programm, das die Daten der Mailbox durchläuft, und wenn Sie eine Zeile finden, die mit `From ` beginnt, teilen Sie die Zeile mit der Funktion split in Wörter auf.
Wir sind daran interessiert, wer die Nachricht gesendet hat, also was das zweite Wort in der `From `-Zeile ist.

Die Funktion <code>extract_senders</code> soll als Ergebnis (also mit `return`) zurückliefern, wie viele E-Mails in der Inbox stehen, also wie viele Zeilen insgesamt mit <code>From&nbsp;</code> beginnen.

Beispieldatenquelle: [http://data.pr4e.org/mbox-short.txt](http://data.pr4e.org/mbox-short.txt)

<!-- grade_id: cell-a44cf395ba6c1d6a, points: 0, grade: false, solution: true -->
```python
def extract_senders(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ebbb23d541e71fc2, points: 1, grade: true, solution: false -->
```python
# Test für die Funktion extract_senders
from nose.tools import assert_equal

assert_equal(27, extract_senders("files/mbox-short.mbox"))
```

**Aufgabe 3:** Schreiben Sie das Programm um, das den Benutzer zur Eingabe einer Liste von Zahlen auffordert und am Ende das Maximum und Minimum der Zahlen ausgibt, wenn der Benutzer `done` eingibt.

Schreiben Sie das Programm so, dass es die Zahlen, die der Benutzer eingibt, in einer Liste speichert und die Funktionen `max()`und `min()` verwendet, um die größte und kleinste Zahl nach Abschluss der Schleife zu ermitteln.

Geben Sie ein Tupel (oder eine Liste) mit 2 Elementen, nämlich dem gesuchten Minimum und Maximum  zurück (ebenfalls mit `return`). Wie bei den Aufgaben 1 und 2 prüft auch hier der Test keine Ausgaben mit `print`, sondern ob die Funktion bestimmte Ergebnisse zurückliefert.

<!-- grade_id: cell-bbeb6a1e0cb55e0e, points: 0, grade: false, solution: true -->
```python
def min_max():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4d6a6440a688b87f, points: 1, grade: true, solution: false -->
```python
# Test für die Funktion min_max

def myinput(inputs):
    def foo(a):
        print(a)
        input.counter += 1
        if input.counter > len(inputs) :
            raise "Nicht genügend inputs"
        else:
            return inputs[input.counter-1]
    foo.counter = 0
    return foo

from nose.tools import assert_equal, assert_not_equal

input_old = input
input = myinput(['2','4', '6', 'done'])
s1 = min_max()
input = input_old

assert_equal(6, s1[0], "Die Ausgabe stimmt nicht")
assert_equal(2, s1[1], "Die Ausgabe stimmt nicht")
```