<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

**Aufgabe 1:** Schreiben Sie ein Programm, dass jede E-Mail-Nachricht danach kategorisiert, an welchem Wochentag sie versendet wurde.

Suchen Sie dazu nach Zeilen, die mit <code>From&nbsp;</code> beginnen, suchen Sie dann nach dem dritten Wort und führen Sie eine laufende Zählung der einzelnen Wochentage durch.
Am Ende des Programms geben Sie Ihr Dictionarys zurück (die Reihenfolge spielt keine Rolle).

<!-- grade_id: cell-5863b0cfcbd0e4b5, points: 0, grade: false, solution: true -->
```python
def count_mails_by_weekday(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9cbd479dbf9b0325, points: 1, grade: true, solution: false -->
```python
# Test für die Funktion count_mails_by_weekday
from nose.tools import assert_equal

fname = "files/mbox-short.mbox"
res = count_mails_by_weekday(fname)

assert_equal(3, len(res), "Es stehen zu viele Werte im dict")
assert_equal(20, res["Fri"], "Die Zahl für Freitage stimmt nicht")
assert_equal(1, res["Sat"], "Die Zahl für Samstag stimmt nicht")
assert_equal(6, res["Thu"], "Die Zahl für Donnerstage stimmt nicht")
```

**Aufgabe 2:** Schreiben Sie ein Programm, um ein E-Mail-Protokoll zu lesen. Erzeugen Sie ein Histogramm mithilfe eines Dictionarys, um zu zählen, wie viele Nachrichten von den einzelnen E-Mail-Adressen gekommen sind. Geben Sie das Dictionary dann zurück.

<!-- grade_id: cell-7442cea4c24e7cff, points: 0, grade: false, solution: true -->
```python
def count_mail_adresses_by_weekday(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6b7600855a664333, points: 1, grade: true, solution: false -->
```python
# Test für die Funktion count_mails_by_weekday
from nose.tools import assert_equal

fname = "files/mbox-short.mbox"
res = count_mail_adresses_by_weekday(fname)

assert_equal(11, len(res), "Es stehen zu viele Werte im dict")
assert_equal({'stephen.marquard@uct.ac.za': 2, 'louis@media.berkeley.edu': 3, 'zqian@umich.edu': 4, 'rjlowe@iupui.edu': 2, 'cwen@iupui.edu': 5, 'gsilver@umich.edu': 3, 'wagnermr@iupui.edu': 1, 'antranig@caret.cam.ac.uk': 1, 'gopal.ramasammycook@gmail.com': 1, 'david.horwitz@uct.ac.za': 4, 'ray@media.berkeley.edu': 1},
             res, "Die gezählten E-Mails stimmen nicht.")
```

**Aufgabe 3:** Bauen Sie ihr Programm so um, dass es die E-Mail Adresse mit den meisten Nachrichten zurück gibt.

Nachdem alle Daten gelesen und das Dictionary erstellt wurde, durchsuchen Sie das Dictionary mithilfe einer Schleife nach dem Maximum (siehe Kapitel 5: Maximum und Minimum ermitteln), um herauszufinden, wer die meisten Nachrichten hat, und berechnen Sie aus, wie viele Nachrichten die Person bekommen hat.

Geben Sie ein Tupel (oder eine Liste) mit 2 Elementen, nämlich dem String der gesuchten E-Mail Adresse und der Anzahl der E-Mails (mit `return`) zurück.

<!-- grade_id: cell-05713fbbe9800a31, points: 0, grade: false, solution: true -->
```python
def max_count_mail_adresses(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-114a763e702f3664, points: 1, grade: true, solution: false -->
```python
# Test für die Funktion  max_count_mails
from nose.tools import assert_equal

fname = "files/mbox-short.mbox"
res = max_count_mail_adresses(fname)

assert_equal("cwen@iupui.edu", res[0], "Die E-Mail Adresse ist falsch")
assert_equal(5, res[1], "Die E-Mail Anzahl ist falsch")
```

**Aufgabe 4:** Schreiben Sie ein Programm, das nur den Domänennamen (anstelle der Adresse) aufzeichnet, von dem die Nachricht gesendet wurde. Es soll also nicht aufgezeichnet werden, von welchem konkreten Absender die E-Mail kam (das wäre die gesamte E-Mail-Adresse). Geben Sie am Ende der Funktion `count_mail_domains` das erstellte Dictionary (mit `return`) zurück.

*Tipp: irgendwas@domäne*

<!-- grade_id: cell-e7bab576cc6e0f06, points: 0, grade: false, solution: true -->
```python
def count_mail_domains(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-532e64bcfc40154e, points: 1, grade: true, solution: false -->
```python
# Test für die Funktion  count_mails_by_domains
from nose.tools import assert_equal
from nose.tools import assert_dict_equal

fname = "files/mbox-short.mbox"
res = count_mail_domains(fname)

assert_equal(6, len(res), "Die Anzahl der Werte ist falsch")
#assert_equal({'uct.ac.za': 6, 'media.berkeley.edu': 4, 'umich.edu': 7, 'iupui.edu': 8, 'caret.cam.ac.uk': 1, 'gmail.com': 1},
#             res, "Die Werte im dict sind falsch")
assert_dict_equal({'uct.ac.za': 6, 'media.berkeley.edu': 4, 'umich.edu': 7, 'iupui.edu': 8, 'caret.cam.ac.uk': 1, 'gmail.com': 1},
             res, "Die Werte im dict sind falsch")
```