<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

**Aufgabe 1:** Das Programm aus Kapitel 9, Übung 3 soll auf folgende Weise überarbeitet werden:

Die <code>From&nbsp;</code>-Zeilen sollen gelesen und geparst werden, um im Folgenden die Adressen aus der Zeile zu extrahieren. Dabei soll mit Hilfe eines Wörterbuchs die Anzahl der Nachrichten von jeder Person gezählt werden.

Nachdem alle Daten gelesen wurden, soll die Person mit den meisten Nachrichten ausgegeben werden, indem eine Liste von (count, email)-Tupeln aus dem Wörterbuch erstellt wird. Dann muss die Liste in umgekehrter Reihenfolge sortiert und schließlich die Person mit den meisten Nachrichten zurückgegeben werden.

Die Funktion `max_count_mail_adresses` soll also ein Tupel zurückliefern (also mit `return`), bei dem das erste Element der String der gesuchten E-Mail Adresse ist und das zweite Element die Anzahl der gesendeten E-Mails.

<!-- grade_id: cell-88a0e8684be45c0a, points: 0, grade: false, solution: true -->
```python
def max_count_mail_adresses(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-871ba8cd03698f18, points: 1, grade: true, solution: false -->
```python
# Test für die Funktion max_count_mails
from nose.tools import assert_equal

fname = "files/mbox-short.mbox"
res = max_count_mail_adresses(fname)

assert_equal("cwen@iupui.edu", res[0], "Die E-Mail Adresse ist falsch")
assert_equal(5, res[1], "Die E-Mail Anzahl ist falsch")
```

**Aufgabe 2:** Schreiben Sie ein Programm, welches die Häufigkeit von Nachrichten pro Tageszeit ermittelt. Die Tageszeit soll in Intervalle von jeweils einer Stunde eingeteilt werden. Es genügt also, den Stunden-Anteil des Zeitstempels zu betrachten.

Die Stunden sollen aus der <code>From&nbsp;</code>-Zeile extrahiert werden, indem die Uhrzeit-Zeichenfolge gefunden und diese Zeichenfolge dann anhand des Doppelpunkts in Teile zergliedert wird. Sobald die Nachrichtenhäufigkeit für jede Stunde gesammelt wurde, sollen diese als Dictionary (mit `return`) zurückgegeben werden.

<!-- grade_id: cell-c0448436a6364f59, points: 0, grade: false, solution: true -->
```python
def count_mail_per_hour(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4ad69e362fb2d53a, points: 1, grade: true, solution: false -->
```python
# Test für die Funktion max_count_mails
from nose.tools import assert_equal

fname = "files/mbox-short.mbox"
res = count_mail_per_hour(fname)

assert_equal(12, len(res), "Die Uhrzeit-Anzahl ist falsch")
```

**Aufgabe 3:** Schreiben Sie ein Programm, das eine Datei liest und die Ziffern in absteigender Reihenfolge der Häufigkeit ausgibt. Das Programm soll nur die Ziffern `0`-`9` zählen. Es soll keine Leerzeichen, Buchstaben, Interpunktionszeichen oder irgendetwas anderes als die Ziffern `0`-`9` zählen.

Die Funktion `count_digits` soll ein Dictionary (mit `return`) zurückgegeben werden, bei dem die Ziffern `0`-`9` als Schlüssel und die entsprechenden Häufigkeiten als Werte eingetragen sind.

*Tipp: Laden Sie die Dateien mit open(filename, 'r', encoding='utf8')*

Beispieltext Deutsch: https://de.wikipedia.org/wiki/Computer

<!-- grade_id: cell-7ed38953021a3d0b, points: 0, grade: false, solution: true -->
```python
def count_digits(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f4ef98edff3a5ccf, points: 1, grade: true, solution: false -->
```python
# Test für die Funktion count_chars
from nose.tools import assert_equal, assert_dict_equal

# Test

wikipedia_computer = '''Charles Babbage und Ada Lovelace gelten durch die von Babbage 1837 entworfene Rechenmaschine Analytical Engine als Vordenker des modernen universell programmierbaren Computers. Konrad Zuse (Z3, 1941 und Z4, 1945) in Berlin, John Presper Eckert und John William Mauchly (ENIAC, 1946) bauten die ersten funktionstüchtigen Geräte dieser Art. Bei der Klassifizierung eines Geräts als universell programmierbarer Computer spielt die Turing-Vollständigkeit eine wesentliche Rolle. Sie ist benannt nach dem englischen Mathematiker Alan Turing, der 1936 das logische Modell der Turingmaschine eingeführt hatte.[1][2]'''
with open('test.txt', 'w') as f:
    f.write(wikipedia_computer)

fname = "test.txt"
res = count_digits(fname)
assert_equal(9, len(res), "Die Anzahl der Ziffern stimmt nicht")
assert_dict_equal({'1': 7, '8': 1, '3': 3, '7': 1, '9': 4, '4': 4, '5': 1, '6': 2, '2': 1},res, "Die gezählten Werte stimmen nicht")
```