<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

# Reguläre Ausdrücke

In diesem Aufgabenblatt wollen wir uns einige Prblemstellungen ansehen, die man mit *regulären Ausdrücken* lösen kann.
Im ersten Teil sollen Sie verschiedene reguläre Ausdrücke schreiben, um Informationen aus einem kurzen Text zu extrahieren.
Im zweiten Teil schauen wir uns ein Beispiel für ein Einsatzgebiet an, in dem reguläre Ausdrücke häufig verwendet werden: dem Abgreifen von Information aus Internetseiten.

*Hinweis:* Beim Finden der rgulären Ausdrücke kann Ihnen die Seite [regex101.com](https://regex101.com) helfen.

Betrachten Sie folgenden kurzen Text über die Fachhochschule Südwestfalen.

```python
text = """Die Fachhochschule Südwestfalen mit Sitz in 58590 Iserlohn besitzt weitere Standorte in
58095 Hagen, 59872 Meschede, 59494 Soest und 58507 Lüdenscheid. Der Rektor der Hochschule ist seit
2008 Claus Schuster. Die Vorgänger im Amt waren Teusner, Michael und J. Liese."""
```

**Aufgabe:** Finden Sie alle Postleitzahlen im Text mit einem regulären Ausdruck. Die Postleitzahlen sollen als Liste (von Strings) unter dem Namen `plzn` abgespeichert werden.

<!-- grade_id: cell-a1d2ad970ab1fa0d, points: 0, grade: false, solution: true -->
```python
import re

plzn = []

# YOUR CODE HERE
raise NotImplementedError()

print(plzn)
```

<!-- grade_id: cell-ab037d4f84fc0314, points: 0, grade: true, solution: false -->
```python
# Test Zelle
assert len(plzn)==5, "Falsch Anzahl von Postleitzahlen gefunden"
for p in [58590, 58095, 59872, 59494, 58507]:
    assert str(p) in plzn, f"PLZ {p} nicht gefunden"
```

In den nächsten Aufgaben wollen wir alle Eigennamen im Text identifizieren. Dazu gehen wir schrittweise vor und verbessern unser Lösung nach und nach.

**Aufgabe:** Schreiben Sie zwei reguläre Ausdrücke um mögliche Eigennamen zu identifizieren. Die Namen sollen nach folgenden Musters zusammengesetzt sein:

- Regulärer Ausdruck 1: `Wort Wort` oder
- Regulärer Ausdruck 2: `Wort, Wort`

wobei `Wort` jeweils mit einem Großbuchstaben beginnen muss, dem eine Reihe von mindestens einem Kleinbuchstaben folgen muss.

Beachten Sie, dass im Deutschen die Buchstaben auch Umlaute und das Eszett enthalten können.

<!-- grade_id: cell-e665bce69c34fd2e, points: 0, grade: false, solution: true -->
```python
namen = []

# YOUR CODE HERE
raise NotImplementedError()

print(namen)
```

<!-- grade_id: cell-bc3e7e390d44b2a5, points: 0, grade: true, solution: false -->
```python
# Test Zelle

assert "Claus Schuster" in namen, "Eigenname 'Claus Schuster' sollte mit einem einfachen re. Ausdruck erkannt werden"
uml = False
for n in namen:
    if 'Vorgänger' in n:
        uml = True
assert uml, "Auch 'Vorgänge' sollte hier erkannt werden"
```

Artikel (z.b. *Der*, *Die* oder *Das*), die nach einem Zeilenumbruch oder Nach einem Punkt folgen, "verunreinigen" leider unsere Ergebnisse. So wird mit einem einfachen regulären Ausdruck etwa der Teilstring *Die Fachhochschule* als Eigenname erkannt.

Um dies zu umgehen, können wir komplexere reguläre Audrücke aufstellen, die mit solchen Sonderfällen umgehen können. Eine Alternative, die oftmals etwas einfacher programmiert werden kann, ist, den ursprünglichen Text vorzuverarbeiten. 

**Aufgabe:** Erzeugen sie aus der Zeichenkette `text` eine neue Zeichenkette in der alle Vorkommen von der bestimmten Artikel *Der*, *Die* und *Das* am <u>Zeilenanfang und Satzanfang</u> entfernt sind. Wenden Sie die regulären Ausdrücke aus der vorherigen Aufgabe erneut an, um nun die Eigennamen zu extrahieren.

Sie sollten nun die Namen *Fachhochschule Südwestfalen*, *Claus Schuster* und *Teusner, Michael* finden.

*Hinweise:*
 - Das Löschen der Zeichen können Sie über die `re`-Funktion `sub()` erledigen
 - Wenn Sie in einem regulären Ausdruck mehrere Möglichkeiten für eine Zeichenfolge angeben, spricht man von Alternativen. Alternativen werden durch einen senkrechten Strich `|` (auch *Pipe*-Symbol genannt) voneinander getrennt. `r'T(ic|ac|oe)'` passt also gleichermaßen auf *Tic*, *Tac* und *Toe*.

<!-- grade_id: cell-5ff1ff87b6cef337, points: 0, grade: false, solution: true -->
```python
namen = []

# YOUR CODE HERE
raise NotImplementedError()

print(namen)

```

<!-- grade_id: cell-51026b1cca4381e8, points: 0, grade: true, solution: false -->
```python
# Test Zelle
assert len(namen)==3, "Falsche Anzahl von Matches"
alle_gefunden = True
for n in ['Fachhochschule Südwestfalen', 'Claus Schuster', 'Teusner, Michael']:
    if n not in namen:
        alle_gefunden = False
assert alle_gefunden, "Nicht alle Namen extrahiert"
```

Manchmal werden bei Personennamen die Vornamen durch ein Initial abgekürzt, wie hier im Beispiel bei Jörg Liese mit *J. Liese*.

**Aufgabe:** Ändern Sie die reguläre Ausdrücke so ab, dass auch abgekürzte Vornamen erkannt werden. Also:
- Regulärer Ausdruck 1: `Wort Wort` oder `W. Wort` oder
- Regulärer Ausdruck 2: `Wort, Wort` oder  `Wort, W.`

wobei `Wort` jeweils mit einem Großbuchstaben beginnen muss, dem eine Reihe von mindestens einem Kleinbuchstaben folgen muss. `W.` sthet für einen Großbuchstaben, aud den direkt ein Punkt folgt.

*Hinweis:* Wenn Sie in einem regulären Ausdruck Alternativen mit einem senkrechten Strich `|` angeben, müsssen sie diese Alternativen häufig mit Klammern gruppieren. `r'A|U|Im'` passt also auf die Zeichenketten *A*, *U* und *Im*. Möchte man *Am*, *Um* und *Im* matchen, müsste der Ausdruck `r'(A|U|I)m'` lauten.

Allerdings werden in regulären Ausdrücken *zusammengefasste Teilausdrücke* (die sogenannten *Capturing Groups*) mit runden Klammern eingefasst. In unserer Ergebniosliste würden also die geklammerten Alternativen als einzelne Teil-Matches erscheinen.

Ein Weg um dieses Bilden von Teilausdrücken zu vermeiden, sind *Non-Capturing Groups*. Im Grunde müssen wir unserer Klammerung sagen, dass der geklammerte Ausdruck nicht in einer Gruppe resultieren soll. In den meisten Sprachen für reguläre Ausdrücke macht man dies, indem man der öffnenden Klammer ein `:?` nachstellt. In unserem Beispiel von oben sieht das folgendermaßen aus: `r'(?:A|U|I)m)'`

<!-- grade_id: cell-6af7fed24481e149, points: 0, grade: false, solution: true -->
```python
namen = []

# YOUR CODE HERE
raise NotImplementedError()

print(namen)
```

<!-- grade_id: cell-9228fafb9303a2c0, points: 0, grade: true, solution: false -->
```python
# Test Zelle
assert len(namen)==4, "Falsche Anzahl von Matches"
alle_gefunden = True
for n in ['Fachhochschule Südwestfalen', 'Claus Schuster', 'J. Liese', 'Teusner, Michael']:
    if n not in namen:
        alle_gefunden = False
assert alle_gefunden, "Nicht alle Namen extrahiert"
```

Zum Schluss wollen wir nun aus den gefunden Namen noch Tupel der Form `(Vorname, Nachname)` bilden. Damit es nicht zu kompliziert wird, lassen wir an dieser Stelle *Fachhochschule* mal als Vornamen, und *Südwestfalen* als Nachnamen durchgehen:-)

**Aufgabe:** Erstellen Sie eine Liste `namen` die alle Eigennamen in Form der oben beschrieben Tupel enthält.

*Hinweis:* Für diese Aufgabe können *Capturing Groups* sehr hilfreich sein. Wir wissen, an welcher Stelle im regulären Ausdruck der Vorname und wo der Nachname steht, können wir beide Teile gruppieren. Ein *Match* `m` hat dann zwei Komponenten, die man mit `m[0]` und `m[1]` einzeln ansprechen und in die richtige Reihenfolge bringen kann.

<!-- grade_id: cell-d95e63f752a996e9, points: 0, grade: false, solution: true -->
```python
namen = []

# YOUR CODE HERE
raise NotImplementedError()

namen
```

<!-- grade_id: cell-913253362f97c562, points: 0, grade: true, solution: false -->
```python
# Test Zelle
assert len(namen)==4, "Falsche Anzahl von Matches"
alle_gefunden = True
tupel = [('Fachhochschule', 'Südwestfalen'), ('Claus', 'Schuster'), ('J.', 'Liese'), ('Michael', 'Teusner')]
for n in tupel:
    if n not in namen:
        alle_gefunden = False
assert alle_gefunden, "Nicht alle Namen extrahiert"
```

## Beliebteste Vornamen bestimmen

Es gibt einige Webseiten auf denen die beliebtesten Vornamen für Neugeborene der letzten Jaher abrufbar sind.
Eine dieser Seiten ist [https://gfds.de/vornamen/beliebteste-vornamen](https://gfds.de/vornamen/beliebteste-vornamen) von der *Gesellschaft für deutsche Sprache e. V.*

Unser Ziel ist es, den HTML-Text dieser Seite abzurufen und die gelisteten Vornamen zu extrahieren. Wir wollen diese Namen mitzählen und am Ende eine Rangliste der beliebtesten Namen ausgeben.

**Aufgabe:** Rufen Sie die URL `https://gfds.de/vornamen/beliebteste-vornamen` auf und lesen Sie den Inhalt der Seite als `UTF-8` kodierten String in die Variable `text`.

<!-- grade_id: cell-1126e46626045b97, points: 0, grade: false, solution: true -->
```python
import urllib
import ssl

# Create SSL Context
ctx = ssl.create_default_context()
ctx.check_hostname = False
ctx.verify_mode = ssl.CERT_NONE

url = 'https://gfds.de/vornamen/beliebteste-vornamen/'

# In text soll der HTML Quelltext der Seite (als String) gespeichert werden
text = ""

# YOUR CODE HERE
raise NotImplementedError()

print(text)
```

<!-- grade_id: cell-456e794396f023f0, points: 0, grade: true, solution: false -->
```python
assert text[:15] == "<!DOCTYPE html>", "Falsches Format."
```

**Aufgabe:** Extrahieren Sie den HTML-Body aus dem gesamten HTML-Quelltext. Schreiben Sie dazu einen entsprechenden *regulären Ausdruck* `regex`. Der Body einer Internetseite ist die HTML-Tags `<body>` `</body>` geklammert.

<!-- grade_id: cell-3fff87940931109d, points: 0, grade: false, solution: true -->
```python
import re

re_options = re.MULTILINE | re.IGNORECASE | re.DOTALL 

regex = r'.*'

# YOUR CODE HERE
raise NotImplementedError()

match = re.search(regex, text, re_options)
if match:
    body = match.group(1)
else:
    print('Nichts gefunden')
```

<!-- grade_id: cell-e9e56265f921a24d, points: 0, grade: true, solution: false -->
```python
assert body!='', "Match ist leer. RegEx stimmt nicht."
assert '<body>' not in body, "HTML body soll extrahiert werden"
```

**Aufgabe:** Schauen Sie in den HTML-Quelltext der Seite. Die Kindernamen tauchen alle mit dem selben Muster auf der Seite auf. Finden Sie alle Vorkommen von Namen auf der Seite und zählen Sie die Vorkommen der einzelnen Namen mit.

Erstellen Sie eine Liste `rangliste` aller Namen geordnete nach der Anzahl ihrer Vorkommen auf der Seite. Die einzelnen Einträge von `rangliste` sollen Tupel aus 2 Elementen sein. Das erste Element ist die Anzahl der Nennungen des Namens, der zweite ist der Name als String), z.B. `(10, 'Tim')`.

Geben Sie die häufigsten 20 Namen aus.

<!-- grade_id: cell-6cf4efddfa8c12d0, points: 0, grade: false, solution: true -->
```python
namen = {}

# YOUR CODE HERE
raise NotImplementedError()

rangliste[:20]
```

<!-- grade_id: cell-e604547fb54e3d19, points: 0, grade: true, solution: false -->
```python
assert rangliste[0][0]==44, "Der Namen 'Alexander' müsste eigentlich 44-mal vorkommen"
assert rangliste[0][1]=='Alexander', "Bei der Rangliste stimmt etwas nicht"
```

**Aufgabe:** Extrahieren Sie alle Links auf HTTP (oder HTTPS) Seiten aus dem Body. Schreiben Sie den regulären Ausdruck so, dass nur der Webseiten Link, aber keine zusätzlichen Artefakte extrahiert werden. In der Liste `links` sollen alle Links der Seite gespeichert werden. Dabei soll `links` keine doppelten Einträge enthalten.

<!-- grade_id: cell-792509d01f91029e, points: 0, grade: false, solution: true -->
```python
links = []

# YOUR CODE HERE
raise NotImplementedError()

for l in links[:20]:
    print(l)
```

<!-- grade_id: cell-b596a8b33dfd1c92, points: 0, grade: true, solution: false -->
```python
assert "https://gfds.de/agbs/" in links, "Nicht alle Links enthalten"
for l in links:
    assert l[:3]=="htt", "Kein HTTP Link"
```


## Zusammenfassung: Reguläre Ausdrücke in Python

Das Modul `re` liefert Methoden, um mit regulären Ausdrücken in Python zu arbeiten.
Eine Suche mit regulären Ausdrücken kann wie folgt durchgeführt werden:
```python
match = re.search(pat, string)
```

Die Methode `re.search()` erhält als Parameter einen regulären Ausdruck `pat` sowie eine Zeichenkette `string` und sucht nach diesem Muster innerhalb des Strings.
Ist  die Suche erfolgreich ist, gibt `search()` ein Match-Objekt zurück, andernfalls `None`.
Nach einer `search`-Operation sollte Sie also normalerweise prüfen, ob die Suche erfolgreich war:

```python
str = 'an example word:cat!!'
match = re.search(r'word:\w\w\w', str)
# If-statement after search() tests if it succeeded
if match:
  print 'found', match.group() ## 'found word:cat'
else:
  print 'did not find'
```


Die Zuweisung `match = re.search()` speichert das Suchergebnis in einer Variablen namens `match`.
Dann testet die `if`-Anweisung die Übereinstimmung -- wenn wahr, war die Suche erfolgreich und `match.group()` ist der passende Text (zB 'word:cat'). Andernfalls, wenn die Übereinstimmung falsch ist (Keine, um genauer zu sein), war die Suche nicht erfolgreich und es gibt keinen übereinstimmenden Text.

Das `r` am Anfang des Muster-Strings bezeichnet einen "rohen" Python-String, der Backslashes ohne Änderung durchläuft, was für reguläre Ausdrücke sehr praktisch ist (Java braucht diese Funktion dringend!). Ich empfehle, dass Sie aus Gewohnheit immer Musterzeichenfolgen mit dem `r` schreiben.

Die Stärke von regulären Ausdrücken besteht darin, dass sie Muster angeben können, nicht nur feste Zeichen. Hier sind die grundlegendsten Muster, die einzelnen Zeichen entsprechen:

- `a`, `X`, `9`, `<` gewöhnliche Zeichen stimmen einfach genau überein. Die Metazeichen, die nicht zu sich selbst passen, weil sie eine besondere Bedeutung haben, sind: `.` `^` `$` `*` `+` `?` `{` `[` `]` `\` `|` `(` `)`
- `[...]` beschreibt eine Zeichenklasse. Alle Zeichen in den eckigen Klammern passen auf das Muster. Die meisten Metazeichen müssen innerhalb einer Zeichenklasse nicht maskiert werden.
- `.` (ein Punkt) entspricht jedem einzelnen Zeichen außer Newline `'\n'`
- `\w` (Kleinbuchstabe w) entspricht einem "Wort"-Zeichen: einem Buchstaben oder einer Ziffer oder einem Unterstrich `[a-zA-Z0-9_]`. 
- `\W` (Großbuchstabe W) entspricht jedem Nicht-Wort-Zeichen.
- `\b` Grenze zwischen Wort und Nicht-Wort
- `\s` (Kleinbuchstabe s) entspricht einem einzelnen Leerzeichen -- Leerzeichen, Zeilenumbruch, Return, Tabulator, Form `[ \n\r\t\f]`.
- `\S` (Großbuchstabe S) entspricht jedem Zeichen, das kein Leerzeichen ist.
- `\t`, `\n`, `\r` Tabulator, Zeilenumbruch, Return
- `\d` Dezimalziffer `[0-9]`
- `^`, `$` Zeilenanfang bzw. -ende
- `\` maskiert Metazeichen. Verwenden Sie also beispielsweise `\.` für einen Punkt oder `\\` für einen Schrägstrich

Eine Suche mit `search` durchläuft die Zeichenfolge von Anfang bis Ende und stoppt bei der ersten gefundenen Übereinstimmung mit dem regulären Ausdruck.
Das gesamte Muster muss übereinstimmen, aber nicht die gesamte Zeichenfolge.

Wenn `match = re.search(pat, str)` erfolgreich, ist `match` nicht `None`. Mit `match.group()` erhält man den als passend gefundenen Textabschnitt.

Mit `+` und `*` können Sie angeben, dass bestimmte Zeichen oder Abschnitte mehrfach (also Widerholt) im passenden Textabschnitt vorkommen dürfen. 
- `+` 1 oder mehr Vorkommen des Musters links davon. Z.B. passt `'i+'` auf *ein oder mehrere i's* 
- `*` 0 oder mehr Vorkommen des Musters links davon.
- `?` Übereinstimmung mit 0 oder 1 Vorkommen des Musters links davon

Die *Quantoren* `+` und `*` arbeiten *gready*, d.h. sie konsumieren so viele Zeichen wie möglich.


Die "Gruppe"-Funktion eines regulären Ausdrucks ermöglicht es, Teile des übereinstimmenden Textes herauszusuchen. Angenommen, wir möchten E-Mail-Adressen finden und dabei den Benutzernamen und die Domäne (also die Zeichen vor bzw. nach dem @-Symbol) extrahieren. 
In einem regulären Ausdrück könnten wir beide Teile klammern `r'([\w.-]+)@([\w.-]+)'`. In diesem Fall ändern die Klammern nicht, was mit dem Muster übereinstimmt, sondern bilden logische "Gruppen" innerhalb des Übereinstimmungstextes. Bei einer erfolgreichen Suche ist `match.group(1)` der Übereinstimmungstext, der der ersten linken Klammer entspricht, und `match.group(2)` ist der Text, der der zweiten linken Klammer entspricht. Ein `match.group()` ohne Index gibt den gesamten Match-Text an.

Einführung zu regulären Ausdrücken in Python:<br>
[1] [Google Education, *Python Regular Expressions*, 2021](https://developers.google.com/edu/python/regular-expressions)