<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

**1.** Importieren Sie das numpy Package mit dem namen ``np``.

<!-- grade_id: cell-2c216d4a55e3848c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**2.** Erstellen Sie in der Funktion ``create_null_vector`` einen ``null`` Vector der größe 10 und geben Sie diesen zurück.

<!-- grade_id: cell-d184b11b2b0aeffa, points: 0, grade: false, solution: true -->
```python
def create_null_vector():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-21527004334d46bd, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(np.zeros(10).all(), create_null_vector().all())
```

**3.** Schreiben Sie eine Funktion ``array_size``, welche die Speichergröße eines übergebenen Vectors berechnet und zurück gibt.

<!-- grade_id: cell-de5ffe690c4f997b, points: 0, grade: false, solution: true -->
```python
def array_size(vec):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-56eee22182c8019c, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

temp_vector = np.zeros((10, 10))
assert_equal(800, array_size(temp_vector))
```

**4.** Schreiben Sie eine Funktion ``create_vector_with_range``, welche einen Vector mit Werten erzeugt. Geben Sie den erzeugten Vector zurück. Als Parameter soll der Start und das Ende des Zahlenbereichs übergeben werden.

<!-- grade_id: cell-a736159a0423c444, points: 0, grade: false, solution: true -->
```python
def create_vector_with_range(start, end):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e82708912f893caa, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(np.arange(10, 50).all(), create_vector_with_range(10, 50).all())
```

**5.** Schreiben Sie eine Funktion ``array_3_3_3`` welche eine dreidimensionale Matrix mit zufälligen Werten erzeugt.

<!-- grade_id: cell-aee00ff400d84998, points: 0, grade: false, solution: true -->
```python
def array_3_3_3():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ad198a08b7430541, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal((3, 3, 3), array_3_3_3().shape, "Die Arraygröße stimmt nicht.")
```

**6.** Schreiben Sie eine Funktion ``calculate_array_extrema`` welche eine Matrix als Parameter übernimmt und ihren höchsten sowie niedrigsten Wert als Tupel zurückgibt.

<!-- grade_id: cell-a9ffdc7376513ba1, points: 0, grade: false, solution: true -->
```python
def calculate_array_extrema(arr):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d727c43cee9a8f83, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

temp = np.arange(21,30).reshape(3, 3)
assert_equal((21, 29), calculate_array_extrema(temp))
```

**7.** Schreiben Sie eine Funktion ``checkboard_matrix`` welche eine 2-Dimensionale Matrix mit Schachbrettmuster erzeugt. Das Schachbrett soll aus den Einträgen ``0`` und ``1`` bestehen. Mit den Parametern ``x`` und ``y`` kann die Matrixgröße übergeben werden.

<!-- grade_id: cell-af25f252752f2f68, points: 0, grade: false, solution: true -->
```python
def checkboard_matrix(x, y):
    # YOUR CODE HERE
    raise NotImplementedError()
checkboard_matrix(3, 5)
```

<!-- grade_id: cell-7fbc83bb06f560f3, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(0, checkboard_matrix(3, 3).min(), "Der niedrigste Wert soll 0 sein.")
assert_equal(1, checkboard_matrix(3, 3).max(), "Der höchste Wert soll 1 sein.")
assert_equal((3, 4), checkboard_matrix(3, 4).shape, "Die Arraygröße stimmt nicht.")
assert_equal(0, checkboard_matrix(3, 4)[1, 1], "Die Schachbrettgenerierung stimmt nicht.")
assert_equal(1, checkboard_matrix(3, 4)[1, 2], "Die Schachbrettgenerierung stimmt nicht.")
```

**8.** Schreiben Sie eine Funktion ``multiply_matrix`` welche zwei übergebene Matrizen multipliziert zurück gibt.

<!-- grade_id: cell-48e9b1f18cd7174e, points: 0, grade: false, solution: true -->
```python
def normalize_matrix(matrix1, matrix2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b56062f21dbdbbc4, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

temp = normalize_matrix(np.ones((5, 3)), np.ones((3, 2)))
assert_equal(3, temp.min(), "Der niedrigste Wert soll 3 sein.")
assert_equal(3, temp.max(), "Der höchste Wert soll 3 sein.")
assert_equal((5, 2), temp.shape)
```

**9** Schreiben Sie eine Funktion ``common_values``, welche die gemeinsamen Werte von zwei übergebenen Matrizen findet und als Standartliste zurückgibt.

<!-- grade_id: cell-51eeef2ec55456ca, points: 0, grade: false, solution: true -->
```python
def common_values(matrix1, matrix2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3913521fd97b9d40, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

m1 = np.arange(9).reshape(3,3)
m2 = np.arange(18).reshape(3, 6)

assert_equal([0, 1, 2, 3, 4, 5, 6, 7, 8], common_values(m1, m2), "Die gefundenen Werte stimmen noch nicht!")
```

**10** Schreiben Sie eine Funktion ``dates`` welche ein Numpyarray mit allen Daten (von Datum) zwischen einem übergebenen Start und Endzeitpunkt generiert. Geben Sie das erstellte Objekt zurück.

*Tipp: dytpe='datetime64[D]'*

<!-- grade_id: cell-f18f69286b6e9637, points: 0, grade: false, solution: true -->
```python
def dates(start, end):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a0d1968cd16f1cb9, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

temp = dates('2022-01', '2022-04')
assert_equal(90, len(temp), "Es wurden nicht entsprechend viele Daten gefunden.")
assert_equal(np.datetime64('2022-01-01'), min(temp), "Die untere Datumsgrenze ist falsch.")
assert_equal(np.datetime64('2022-03-31'), max(temp), "Die obere Datumsgrenze ist falsch.")
```

**11** Schreiben Sie eine Funktion ``small_vector`` welche einen Vektor der größe 10 erstellt und dessen Werte nur den Wertebereich ``0 < x < 1`` besitzen. Geben Sie den erstellten Vektor zurück.

<!-- grade_id: cell-011369ef405566ff, points: 0, grade: false, solution: true -->
```python
def small_vector():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7887c51d75257bde, points: 1, grade: true, solution: false -->
```python
temp = small_vector()
assert min(temp) > 0, "Der kleinste Wert muss größer als 0 sein."
assert max(temp) < 1, "Der größte Wert muss kleiner als 1 sein."
assert 10 == len(temp), "Der Vektor soll die Größe 10 besitzen."
```

**12** Schreiben Sie eine Funktion ``convert_int_to_float`` welche ein übergebenes float Array in ein integer Array überführt und zurückgibt.

<!-- grade_id: cell-a8e4e28d0c26929c, points: 0, grade: false, solution: true -->
```python
def convert_int_to_float(m1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ba2c6b851baf3860, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

temp = (np.random.rand(10)*100).astype(np.float32)
temp2 = convert_int_to_float(temp)
assert_equal(np.int32, type(temp2[1]), "Die Konvertierung war nicht erfolgreich.")
```

**13.** Schreiben Sie ein Programm ``random_matrix`` womit Sie eine 5x3 Matrix erzeugen, welche mit 0en initialisiert ist. Platzieren Sie anschließend in dieser Matrix die übergebene Anzahl von ``1`` an zufälligen Stellen. Geben Sie die generierte Matrix zurück.

<!-- grade_id: cell-523f142a5e9ef09e, points: 0, grade: false, solution: true -->
```python
def random_matrix(number):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d036bf1c547ced30, points: 1, grade: true, solution: false -->
```python
import collections
from nose.tools import assert_equal

temp = random_matrix(4)
unique, counts = np.unique(temp, return_counts=True)
x = dict(zip(unique, counts))

assert_equal(x[1], 4, "Die generierte Anzahl stimmt nicht.")

temp2 = random_matrix(4)
assert_equal(False, (temp==temp2).all(), "Generieren Sie die Wertpositionen zufällig.")
```

**14.** Sortieren Sie das Array nach der Nten Spalte. Implementieren Sie dafür eine Funktion ``sort_matrix``. Die Funktion nimmt als Parameter ein Array und eine Spaltennummer nach welcher sortiert werden soll an. Geben Sie das sortierte Array zurück.

<!-- grade_id: cell-ec64b9fa891cd711, points: 0, grade: false, solution: true -->
```python
def sort_matrix(array, column):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-da8a71f89e9cd320, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
array = np.array([[1,3,5,7,9], [15, 2, 9, 6, 3], [2,4,6,8,10]])

r = sort_matrix(array, 1)
r1 = np.array([[15,2,9,6,3],[1,3,5,7,9],[2,4,6,8,10]])

assert_equal(True, (r==r1).all(), "Das Array wurde nicht richtig sortiert.")
```

**15.** Schreiben Sie eine Funktion ``find_next`` welche ein Array entgegennimmt. Geben Sie den am nächsten an einem übergebenen Wert liegenden Wert zurück. 

<!-- grade_id: cell-c04dc1112e92696d, points: 0, grade: false, solution: true -->
```python
def find_next(array, value):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0e39d64d1728c8bd, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

array = np.array([[1,3,5,7,9], [15, 2, 9, 6, 3], [2,4,6,8,10]])
assert_equal(15, find_next(array, 20), "Es wurde der falsche Wert gefunden.")
```

**16.** Schreiben Sie eine Funktion ``split_array`` die aus einem übergebenen Array die höchsten 5 Zahlen findet und zurückgibt.

<!-- grade_id: cell-93fdc67224c5c399, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-92aa3223fe8d7ce8, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
array = np.arange(40,60,dtype=np.uint32)

assert_equal(True, (np.array([55,56,57,58,59])==split_array(array)).all(), "Es wurden die falschen Elemente gefunden.")
```

**17.** Schreiben Sie eine Funktion ``remove_equal`` welche Zeilen aus einer übergebenen Matrix entfernt, die ausschließlich den selben Wert enthalten (z.B. [1,1,1]). Geben Sie die korrigierte Matrix zurück.

<!-- grade_id: cell-ab03147bf79098c9, points: 0, grade: false, solution: true -->
```python
def remove_equal(array):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-10ec3fb98721646c, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

array = np.array([[4,1,0],[4,2,2],[1,1,3],[3,2,0],[4,4,3],[0,4,2],[4,4,4],[0,1,2],[3,4,2],[3,0,3]])
r = np.array([[4,1,0],[4,2,2],[1,1,3],[3,2,0],[4,4,3],[0,4,2],[0,1,2],[3,4,2],[3,0,3]])

assert_equal(True, (r == remove_equal(array)).all(), "Es wurden nicht die richtigen Werte entfernt.")
```