<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

**Übung 1:**<br>

Der offene Webservice zippopotam.us stellt Informationen zu Postleitzahlen in einer vielzahl von Ländern im JSON-Format zur Verfügung.

Die Route um Daten abzurufen ist wie folgt aufgebaut:
*https://api.zippopotam.us/[country]/[zipcode]"*

Schreiben Sie ein Programm, dass die Informationen zur Postleitzahl `58644` abruft. Zudem soll es die `longitude` sowie die `latitude` in einem Tupel des ersten gefundenen Ortes zurück geben. Die erwarteten Ergebnisse sollen in einem Tupel zurückgegeben werden.<br>

*Tipp: Sie können die URL zum Webservice auch mit dem Browser aufrufen*<br>
*Tipp: Der Country-Code für Deutschland lautet: de*<br>
*Tipp: Interpretieren Sie die JSON Daten mit json.load([daten])*

<!-- grade_id: cell-241baf5527f685bd, points: 0, grade: false, solution: true -->
```python
import json
import urllib

def fetch_58644():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3bd8fa3409d1d3d2, points: 1, grade: true, solution: false -->
```python
"""Tests fetch_58644"""
from nose.tools import assert_equal

assert_equal(('7.6968', '51.3246'), fetch_58644(), "Es wurden die falschen Daten extrahiert.")
```

**Übung 2:**<br>

Bauen Sie das Programm aus Übung 1 so um, dass mit zwei Parametern das Land und die Postleitzahl übergeben werden können.

<!-- grade_id: cell-9c1d4edc9414b33c, points: 0, grade: false, solution: true -->
```python
import json
import urllib

def fetch(country, zipcode):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-84f7900cb7afce02, points: 1, grade: true, solution: false -->
```python
"""Tests fetch"""
from nose.tools import assert_equal

assert_equal(('7.6968', '51.3246'), fetch("de", "58644"), "Es wurden die falschen Daten extrahiert.")
assert_equal(('6.6971', '46.5534'), fetch("ch", "1000"), "Es wurden die falschen Daten extrahiert.")
assert_equal(('28.2294', '-25.7069'), fetch("za", "0002"), "Es wurden die falschen Daten extrahiert.")
```