<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

**Übung 1:**<br>
Schreiben Sie eine Klasse `Standort`. Die Klasse soll die Daten *Postleitzahl* (`zipcode`), *Längengrad* (`longitude`) und Breitengrad (`latitude`) speichern. <br> Zudem soll die Klasse eine Funktion `getLocation` enthalten, welche die `latitude` und die `longitude` als Tupel zurückgibt.<br>
Der Konstruktor von Standort soll die Parameter `zipcode`, `longitude` und `latitude` entgegennehmen und in den entsprechenden Eigenschaften ablegen.

<!-- grade_id: cell-e6d754d6e230ab48, points: 0, grade: false, solution: true -->
```python
class Standort:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ce79c3192186d96e, points: 1, grade: true, solution: false -->
```python
"""test class standort"""
from nose.tools import assert_equal

s = Standort("1111", 1.0, 1.0)
assert_equal((1.0, 1.0), s.getLocation(), "Da stimmt etwas mit der Zuweisung nicht!")

```

Der offene Webservice zippopotam.us stellt Informationen zu Postleitzahlen in einer vielzahl von Ländern im JSON-Format zur Verfügung.

Die Route um Daten abzurufen ist wie folgt aufgebaut:
*https://api.zippopotam.us/[country]/[zipcode]"*


**Übung 2:** <br>

Bauen Sie das Programm aus der <u>Übung 2 von Arbeitsblatt `webservices.ipynb`</u> so um, dass ein <u>Objekt vom Typ Standort</u> mit den erhaltenen Daten zurückgegeben wird.

*Tipp: Der Country-Code für Deutschland lautet: de*<br>
*Tipp: Interpretieren Sie die JSON Daten mit json.load([daten])*

<!-- grade_id: cell-e7aab6c95906bfb5, points: 0, grade: false, solution: true -->
```python
import json
import urllib

def fetch(country, zipcode):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0a5d4c0199321ce1, points: 1, grade: true, solution: false -->
```python
"""Tests fetch"""
from nose.tools import assert_equal

s = fetch("de", "58644")

assert isinstance(s,Standort), "Die Funktion 'fetch' soll ein Objekt vom Typ 'Standort' zurückliefern"
assert_equal(('7.6968', '51.3246'), s.getLocation(), "Es wurden die falschen Daten extrahiert.")
```